```python
# Initialisation Cell
# You should always put imported modules here
import numpy as np
import numpy.testing as nt
import numpy.linalg as LA
from matplotlib import pyplot as plt
np.set_printoptions(suppress=True, precision=7)
```

# CDES Honours - Lab 1


## Instructions

* Read all the instructions carefully.
* Do not rename the notebook, simply answer the questions and resubmit the file to Moodle.
* **Numpy** has a help file for every function if you get stuck. See: https://docs.scipy.org/doc/numpy-1.15.4/reference/
* See these useful links:
    * https://docs.scipy.org/doc/numpy/user/numpy-for-matlab-users.html
    * https://docs.scipy.org/doc/numpy/user/quickstart.html
* **Numpy** is not always required.
* There are also numerous sources available on the internet, Google is your friend!

# Warm-up Exercises

Complete the following warm-up tasks without the use of numpy.

## Question 1

Write a function which takes in a matrix and computes the absolute difference between the sums of its diagonals (i.e. left-to-right and right-to-left).

<!-- grade_id: cell-3ac1b7539439cad0, points: 0, grade: false, solution: true -->
```python
def diagonalDifference(arr):
    # YOUR CODE HERE
    raise NotImplementedError()
    

```

<!-- grade_id: cell-e6927b4b248aa48d, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
test =  [[11, 2,   4],
         [4,  5,   6],
         [10, 8, -12]]

assert(diagonalDifference(test) == 15)
print('Test case passed!!!')
```

<!-- grade_id: cell-290d5589a75e5222, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

# Question 2

Write a function that takes as input an array of integers. The function should compute the fractions of its elements that are positive, negative and are zeros. These fractions should be returned rounded to 4 decimal places.

<!-- grade_id: cell-57cfe1c93be8c7c4, points: 0, grade: false, solution: true -->
```python
def plusMinus(arr):
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-aa8ec8e474837f88, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
arr = [-4, 3, -9, 0, 4, 1]
nt.assert_array_almost_equal(plusMinus(arr),(0.5, 0.3333, 0.1667))
print('Test case passed!!!')
```

<!-- grade_id: cell-b407ab91621a8e55, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

# Question 3

Write a function that takes as input a matrix and determines whether or not the matrix is symmetric. The function should return the strings `not symmetric` or `symmetric`. You may not use any obvious builtin function.

<!-- grade_id: cell-7dc873ee1e1ae80f, points: 0, grade: false, solution: true -->
```python
# Question 1
def symCheck(A):
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-f50586567f175f23, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A = np.array([[1, 2, 3], [2, 1, 3]])
assert(symCheck(A) == 'not symmetric')
print('Test case passed!!!')
```

<!-- grade_id: cell-1a748bc0acde524d, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A = np.array([[1, 2, 3], [2, 1, 3],  [2, 1, 3]])
assert(symCheck(A) == 'not symmetric')
print('Test case passed!!!')
```

<!-- grade_id: cell-e46c86ccd91521b3, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A = np.array([[1, 2, 3], [2, 1, 3],  [3, 3, 3]])
assert(symCheck(A) == 'symmetric')
print('Test case passed!!!')
```

<!-- grade_id: cell-915baceaa05a0dbb, points: 3, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 3 marks
```

<!-- grade_id: cell-d2c1d23eded8be04, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

# Question 4

Write a function that takes in 3 inputs, a list of integers `a`, a integer `k`, and input `q` (which may be an integer or list). The function should rotate the array `a` `k` times. For example, if `ar =  [1, 2, 3]` and `k = 2`, then we rotate the array rightwards, resulting in `[2, 3, 1]`. THe function should then return the `q` indices of `ar`. Thus, in our example, in `q = [1, 2]`, then the function should return the array `[3, 1]`. 

<!-- grade_id: cell-3052d86ef80162af, points: 0, grade: false, solution: true -->
```python
def indexRotation(a, k, q):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4731f7efad11bdee, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
ar   = [1, 2, 3]
k    = 2
q    = [0, 1, 2] 
tans = [2, 3, 1]
assert(tans == indexRotation(ar, k, q))
print('Test case passed!!!')
```

<!-- grade_id: cell-66d317fdef2c8a9e, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
ar   = [1, 2, 3]
k    = 4
q    = [0, 1, 2] 
tans = [3, 1, 2]
assert(tans == indexRotation(ar, k, q))
print('Test case passed!!!')
```

<!-- grade_id: cell-b1910c675e531b32, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

## Main Exercises

Complete the following problems:

## Question 1

Write a function implements a finite difference scheme to solve the heat equation. Specifically, the function should implement an implicit scheme which is backward difference in time and central difference in space. It should take as inputs a time step `dt`, a spatial step `dx`, a number of iterations to march forward in time `N`, the coefficient `D`, an initial function `f(x)` passed as a handle and boundary values `alpha` and `beta`. The function should output the solution space matrix `u`. Recall the heat equation problem as:
$$
u_{t} = Du_{xx},\ \ \quad u(0, t) = \alpha, \ \ u(1, t) = \beta,\ \ \qquad u(x, 0) = f(x) 
$$

Note: the solution matrix here is matrix of all $x$ through time $t$. That is row one is at time zero and the matrix should have `N + 1` rows.

<!-- grade_id: cell-b274afb8bbf0527d, points: 0, grade: false, solution: true -->
```python
def heat_eq_matrix(dt, dx, N, f, D, alpha, beta):
    # YOUR CODE HERE
    raise NotImplementedError()


```

<!-- grade_id: cell-d1e357f3a78c090b, points: 2, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 2 mark
# Unit test
D     = 1
dx    = 0.05
dt    = (dx**2)/(4*D)
alpha = 0
beta  = 0
f     = lambda x: np.exp(-100*(x - 0.5)**2) 
N     = 30
ans   = heat_eq_matrix(dt, dx, N, f, D, alpha, beta)
t1    = np.array([0.       , 0.       , 0.0000001, 0.0000048, 0.0001234, 0.0019305,
       0.0183156, 0.1053992, 0.3678794, 0.7788008, 1.       , 0.7788008,
       0.3678794, 0.1053992, 0.0183156, 0.0019305, 0.0001234, 0.0000048,
       0.0000001, 0.       , 0.       ])
t2    = np.array([0.       , 0.0209475, 0.0452638, 0.0757908, 0.1142184, 0.1603843,
       0.211701 , 0.26306  , 0.3075325, 0.3379146, 0.3487257, 0.3379146,
       0.3075325, 0.26306  , 0.211701 , 0.1603843, 0.1142184, 0.0757908,
       0.0452638, 0.0209475, 0.       ])

nt.assert_array_almost_equal(ans[0, ::], t1, 5)
nt.assert_array_almost_equal(ans[-1, ::], t2, 5)
print('Test case passed!!!')
```

<!-- grade_id: cell-2451db0239faa657, points: 13, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 13 marks
```

## Question 2

Write a function implements a finite difference scheme to solve the heat equation with Neumann boundary conditions. Specifically, the function should implement an explicit scheme which is forward difference in time and central difference in space. It should take as inputs a time step `dt`, a spatial step `dx`, a number of iterations to march forward in time `N`, the coefficient `D`, an initial function `f(x)` passed as a handle and boundary values `alpha` and `beta`. The function should output the solution space matrix `u`. Recall the heat equation problem as:
$$
u_{t} = Du_{xx},\ \ \quad u_x(0, t) = \alpha, \ \ u_x(1, t) = \beta,\ \ \qquad u(x, 0) = f(x) 
$$

Note: the solution matrix here is matrix of all $x$ through time $t$. That is row one is at time zero and the matrix should have `N + 1` rows.

<!-- grade_id: cell-68752a06d6be3520, points: 0, grade: false, solution: true -->
```python
def heat_eq_matrix(dt, dx, N, f, D, alpha, beta):
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-198828b23835c862, points: 2, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 2 mark
# Unit test
dx    = 0.05
D     = 1 
dt    = (dx**2)/(4*D)
alpha = -0.5
beta  = -0.5
f     = lambda x: x*(1 - x)*(x - 0.5)
N     = 50
ans   = heat_eq_matrix(dt, dx, N, f, D, alpha, beta)
t1    = np.array([-0.      , -0.021375, -0.036   , -0.044625, -0.048   , -0.046875,
       -0.042   , -0.034125, -0.024   , -0.012375,  0.      ,  0.012375,
        0.024   ,  0.034125,  0.042   ,  0.046875,  0.048   ,  0.044625,
        0.036   ,  0.021375,  0.      ])
t2    = np.array([ 0.0688784,  0.0461262,  0.0278104,  0.0137554,  0.0036763,
       -0.002811 , -0.0061764, -0.0069613, -0.0057633, -0.003221 ,
        0.       ,  0.003221 ,  0.0057633,  0.0069613,  0.0061764,
        0.002811 , -0.0036763, -0.0137554, -0.0278104, -0.0461262,
       -0.0688784])
nt.assert_array_almost_equal(ans[0, ::], t1, 5)
nt.assert_array_almost_equal(ans[-1, ::], t2, 5)
print('Test case passed!!!')
```

<!-- grade_id: cell-82920529b11ad1b6, points: 13, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 13 marks
```