```python
# Initialisation Cell
# You should always put imported modules here
from math import *
import numpy as np
import numpy.testing as nt
import numpy.linalg as LA
from matplotlib import pyplot as plt
np.set_printoptions(suppress=True, precision=7)
```

# CDES Honours - Test 1 Deferred


## Instructions

* Read all the instructions carefully.
* The test consists of **50 Marks**, with **two** hours available.
* The written section is to be answered in the book provided.
* You must only access Moodle tests and NOT Moodle.
* The programming section is to be answered within this Jupyter notebook and resubmitted to Moodle.
* Do not rename the notebook, simply answer the questions and resubmit the file to Moodle.
* The moodle submission link will expire at exactly **15:00** and **NO** late submission will be accepted. Please make sure you submit timeously!
* You may use the **Numpy** help file at: https://docs.scipy.org/doc/numpy-1.15.4/reference/
* **NB!!!** Anyone caught using Moodle (and its file storing), flash drives or external notes will be awarded zero and reported to the legal office.

## Written Section

* Answer the following questions in the answer book provided.


### Question 1 (10 Marks)

Consider the following difference scheme and conduct a von Neumann stability analysis to determine its stability:
$$
\dfrac{u_n^m - u_n^{m-1}}{\Delta t} = \dfrac{a}{2\Delta x^2}\left(u_{n + 1}^{m} - 2u_{n}^{m} + u_{n - 1}^{m}\right) + \dfrac{a}{2\Delta x^2}\left(u_{n + 1}^{m - 1} - 2u_{n}^{m - 1} + u_{n - 1}^{m - 1}\right)
$$
You may assume that $a>0$.


### Question 2 (5 Marks)

Establish the Taylor expansion for the following finite difference approximation to $u_x$ at the point $(i, j)$ for a uniform mesh:
$$
\frac{\partial u}{\partial x} \approx \frac{-3 u_{i, j}+4 u_{i, j+1}-u_{i, j+2}}{2 \Delta x}
$$

### Question 3 (5 Marks)

Verify that:
$$
\left.\frac{\partial^{3} u}{\partial x^{3}}\right|_{i, j}=\frac{\Delta_{x}^{3} u_{i, j}}{(\Delta x)^{3}}+O(\Delta x)
$$
(Hint: Taylor series)

## Programming Section 

### Question 1 (10 Marks)


Write a function which implements a finite difference scheme to solve the partial differential equation below. Specifically, the function should implement an explicit scheme which is forward difference in time and central difference in space. It should take as inputs a time step `dt`, a spatial step `dx`, a number of iterations to march forward in time `N`, the coefficients `D` and `nu`, an initial function `ic` passed as a handle and boundary values `alpha` and `beta`. The function should output the solution space matrix `u`, as well as the `xvals` (i.e. vector of $x$ steps) and `tvals` (i.e the vector of time steps). The PDE is given below:


$$v_t + \nu v_x = D v_{xx},$$

$$\text{BC}: v(0, t) = \alpha = v(1, t) = \beta = 0;\quad\text{IC}: v(x, 0) = \sin(8\pi x); \quad t \geq 0, \ x \in [0, 1]$$

<!-- grade_id: cell-9e6c20fef606ad4f, points: 0, grade: false, solution: true -->
```python
def heatEqn(dt, dx, N, ic, D, alpha, beta, nu):
    # YOUR CODE HERE
    raise NotImplementedError()
    

```

<!-- grade_id: cell-61884124e5db0039, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
dx = 0.01
dt = 0.01
N  = 2
D  = 0.1
nu = 1
alpha = 0
beta  = 0
ic = lambda x: np.sin(8*np.pi*x) 
u, xx, tt = heatEqn(dt, dx, N, ic, D, alpha, beta, nu)
tans = np.array(  [ 0.       ,  0.2486899,  0.4817537,  0.6845471,  0.8443279,
                    0.9510565,  0.9980267,  0.9822873,  0.9048271,  0.7705132,
                    0.5877853,  0.3681246,  0.1253332, -0.1253332, -0.3681246,
                   -0.5877853, -0.7705132, -0.9048271, -0.9822873, -0.9980267,
                   -0.9510565, -0.8443279, -0.6845471, -0.4817537, -0.2486899,
                   -0.       ,  0.2486899,  0.4817537,  0.6845471,  0.8443279,
                    0.9510565,  0.9980267,  0.9822873,  0.9048271,  0.7705132,
                    0.5877853,  0.3681246,  0.1253332, -0.1253332, -0.3681246,
                   -0.5877853, -0.7705132, -0.9048271, -0.9822873, -0.9980267,
                   -0.9510565, -0.8443279, -0.6845471, -0.4817537, -0.2486899,
                   -0.       ,  0.2486899,  0.4817537,  0.6845471,  0.8443279,
                    0.9510565,  0.9980267,  0.9822873,  0.9048271,  0.7705132,
                    0.5877853,  0.3681246,  0.1253332, -0.1253332, -0.3681246,
                   -0.5877853, -0.7705132, -0.9048271, -0.9822873, -0.9980267,
                   -0.9510565, -0.8443279, -0.6845471, -0.4817537, -0.2486899,
                   -0.       ,  0.2486899,  0.4817537,  0.6845471,  0.8443279,
                    0.9510565,  0.9980267,  0.9822873,  0.9048271,  0.7705132,
                    0.5877853,  0.3681246,  0.1253332, -0.1253332, -0.3681246,
                   -0.5877853, -0.7705132, -0.9048271, -0.9822873, -0.9980267,
                   -0.9510565, -0.8443279, -0.6845471, -0.4817537, -0.2486899,
                    0.       ])
nt.assert_array_almost_equal(tans, u[0])
print('Test case passed!!!')
```

<!-- grade_id: cell-daa46aed8cf80257, points: 9, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 9 marks
```

### Question 2 (10 Marks)

Write a function which implements a finite difference scheme to solve the partial differential equation given in Question 1 (above). Now, however, the function should implement an implicit scheme which is forward in time and central in space.

<!-- grade_id: cell-cae972cea5502f8c, points: 0, grade: false, solution: true -->
```python
def heatEqnImp(dt, dx, N, ic, D, alpha, beta, nu):
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-ecdd711747a7d4d5, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
dx = 0.01
dt = 0.01
N  = 2
D  = 0.1
nu = 1
alpha = 0
beta  = 0
ic = lambda x: np.sin(8*np.pi*x) 
u, xx, tt = heatEqnImp(dt, dx, N, ic, D, alpha, beta, nu)
tans = np.array(  [ 0.       ,  0.2486899,  0.4817537,  0.6845471,  0.8443279,
                    0.9510565,  0.9980267,  0.9822873,  0.9048271,  0.7705132,
                    0.5877853,  0.3681246,  0.1253332, -0.1253332, -0.3681246,
                   -0.5877853, -0.7705132, -0.9048271, -0.9822873, -0.9980267,
                   -0.9510565, -0.8443279, -0.6845471, -0.4817537, -0.2486899,
                   -0.       ,  0.2486899,  0.4817537,  0.6845471,  0.8443279,
                    0.9510565,  0.9980267,  0.9822873,  0.9048271,  0.7705132,
                    0.5877853,  0.3681246,  0.1253332, -0.1253332, -0.3681246,
                   -0.5877853, -0.7705132, -0.9048271, -0.9822873, -0.9980267,
                   -0.9510565, -0.8443279, -0.6845471, -0.4817537, -0.2486899,
                   -0.       ,  0.2486899,  0.4817537,  0.6845471,  0.8443279,
                    0.9510565,  0.9980267,  0.9822873,  0.9048271,  0.7705132,
                    0.5877853,  0.3681246,  0.1253332, -0.1253332, -0.3681246,
                   -0.5877853, -0.7705132, -0.9048271, -0.9822873, -0.9980267,
                   -0.9510565, -0.8443279, -0.6845471, -0.4817537, -0.2486899,
                   -0.       ,  0.2486899,  0.4817537,  0.6845471,  0.8443279,
                    0.9510565,  0.9980267,  0.9822873,  0.9048271,  0.7705132,
                    0.5877853,  0.3681246,  0.1253332, -0.1253332, -0.3681246,
                   -0.5877853, -0.7705132, -0.9048271, -0.9822873, -0.9980267,
                   -0.9510565, -0.8443279, -0.6845471, -0.4817537, -0.2486899,
                    0.       ])
nt.assert_array_almost_equal(tans, u[0])
print('Test case passed!!!')
```

<!-- grade_id: cell-ec3e6aeaa0f32805, points: 9, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 9 marks
```

### Question 3 (5 Marks)


Given the input below, plot your solutions for both schemes on two different sets of axes. That is a 2D plot, illustrating `N` iterations per plot.

```python
dx    = 0.01
dt    = 0.01
N     = 2
D     = 0.1
nu    = 1
alpha = 0
beta  = 0
ic    = lambda x: np.sin(8*np.pi*x) 
```

<!-- grade_id: cell-e3f7296873429171, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Question 4 (5 Marks)

Write a function implements a finite difference scheme to solve the heat equation with Neumann boundary conditions. Specifically, the function should implement an explicit scheme which is forward difference in time and central difference in space. It should take as inputs a time step `dt`, a spatial step `dx`, a number of iterations to march forward in time `N`, the coefficient `D`, an initial function `f(x)` passed as a handle and boundary values `alpha` and `beta`. The function should output the solution space matrix `u`. Recall the heat equation problem as:
$$
u_{t} = Du_{xx},\ \ \quad u_x(0, t) = \alpha, \ \ u_x(1, t) = \beta,\ \ \qquad u(x, 0) = f(x) 
$$

Note: the solution matrix here is matrix of all $x$ through time $t$. That is row one is at time zero and the matrix should have `N + 1` rows.

<!-- grade_id: cell-eabcb39964951637, points: 0, grade: false, solution: true -->
```python
def heat_eq_matrix(dt, dx, N, f, D, alpha, beta):
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-d44e89d66c0319c0, points: 2, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 2 mark
# Unit test
dx    = 0.05
D     = 1 
dt    = (dx**2)/(4*D)
alpha = -0.5
beta  = -0.5
f     = lambda x: x*(1 - x)*(x - 0.5)
N     = 50
ans   = heat_eq_matrix(dt, dx, N, f, D, alpha, beta)
t1    = np.array([-0.      , -0.021375, -0.036   , -0.044625, -0.048   , -0.046875,
       -0.042   , -0.034125, -0.024   , -0.012375,  0.      ,  0.012375,
        0.024   ,  0.034125,  0.042   ,  0.046875,  0.048   ,  0.044625,
        0.036   ,  0.021375,  0.      ])
t2    = np.array([ 0.0688784,  0.0461262,  0.0278104,  0.0137554,  0.0036763,
       -0.002811 , -0.0061764, -0.0069613, -0.0057633, -0.003221 ,
        0.       ,  0.003221 ,  0.0057633,  0.0069613,  0.0061764,
        0.002811 , -0.0036763, -0.0137554, -0.0278104, -0.0461262,
       -0.0688784])
nt.assert_array_almost_equal(ans[0, ::], t1, 5)
nt.assert_array_almost_equal(ans[-1, ::], t2, 5)
print('Test case passed!!!')
```

<!-- grade_id: cell-80f47e1925d3ba03, points: 3, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 3 marks
```