```python
# Initialisation Cell
# You should always put imported modules here
from math import *
import numpy as np
import numpy.testing as nt
import numpy.linalg as LA
from matplotlib import pyplot as plt
import seaborn as sns
#sns.set_style('dark', {'figure.facecolor': '(0,0,0,0)'}) 
sns.set(style="darkgrid")
plt.rcParams['figure.figsize'] = [8, 5]
plt.rcParams['figure.dpi'] = 100
np.set_printoptions(suppress=True, precision=7)
```

# CDES Honours - Lab 3


## Instructions

* Read all the instructions carefully.
* Do not rename the notebook, simply answer the questions and resubmit the file to Moodle.
* **Numpy** has a help file for every function if you get stuck. See: https://docs.scipy.org/doc/numpy-1.15.4/reference/
* See these useful links:
    * https://docs.scipy.org/doc/numpy/user/numpy-for-matlab-users.html
    * https://docs.scipy.org/doc/numpy/user/quickstart.html
* **Numpy** is not always required.
* There are also numerous sources available on the internet, Google is your friend!

## Warmup Exercises

Complete the following problems:

## Question 1

Write a function that takes in a integer `n`. The function should determine the number of divisors occuring within the integer. For example, if in the integer was `12`, then both `1` and `2` are perfect divisors of `12` and thus the function should return `2`.


<!-- grade_id: cell-d71c439ecc75f74b, points: 0, grade: false, solution: true -->
```python
def intDivsor(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f77e8f65cb9806ae, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
test = 12
assert(intDivsor(test) == 2)
print('Test case passed!!!')
```

<!-- grade_id: cell-d141059f1b617cb1, points: 2, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 2 marks
```

<!-- grade_id: cell-a0e655bf715aa805, points: 2, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 2 marks
```

## Question 2

Write a function that takes in a list of integers, $p(1), p(2), \ldots, p(n)$, where each element is distinct and satisfies $1 \leq p(x) \leq n$. For each $x$ where $1 \leq x \leq n$. The function should find the integer $y$ such that $p(p(y)) = x$ and return the appropriate list. For example, given `p = [5, 2, 1, 3, 4]`, each value of $x$ between 1 and 5 is found as follows:

1. $x = 1 = p(3)$, $p(4) = 3$, $\Rightarrow p(p(4)) = 1$
2. $x = 2 = p(2)$, $p(2) = 2$, $\Rightarrow p(p(4)) = 2$
3. $x = 3 = p(4)$, $p(5) = 4$, $\Rightarrow p(p(4)) = 3$
4. $x = 4 = p(5)$, $p(1) = 5$, $\Rightarrow p(p(4)) = 4$
5. $x = 5 = p(1)$, $p(3) = 1$, $\Rightarrow p(p(4)) = 5$

Therefore, the function should return the list `[4, 2, 5, 1, 3]`.

<!-- grade_id: cell-11aab327406194a5, points: 0, grade: false, solution: true -->
```python
def sequenceExplore(p):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-124de6e2b4e75bdb, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
test = [5, 2, 1, 3, 4]
assert(sequenceExplore(test) == [4, 2, 5, 1, 3])
print('Test case passed!!!')
```

<!-- grade_id: cell-55074958c9088317, points: 2, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 2 marks
```

<!-- grade_id: cell-6a982265ae4dea37, points: 2, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 2 marks
```

## Question 2

Write a function that calculates the growth of a tree over time. Specifically, this species of tree goes through 2 cycles of growth per year. Each spring the tree doubles in height, and in winter only grows 1 meter. Given that a tree is planted in spring with a starting height of 1 meter, compute the height of the tree after `n` growth cycles.


<!-- grade_id: cell-ec4e52698980d30a, points: 0, grade: false, solution: true -->
```python
def treeScaling(n):
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-9254fa3f3ec1cfeb, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
test = 5
assert(treeScaling(test) == 14)
print('Test case passed!!!')
```

<!-- grade_id: cell-77786a5c21b08b06, points: 2, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 2 marks
```

<!-- grade_id: cell-c413ce12594e5082, points: 2, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 2 marks
```

## Main Exercises

Complete the following problems:

## Question 1

Write a function implements a finite difference scheme to solve the partial differetial equation below. Specifically, the function should implement a Crank-Nicolson scheme. It should take as inputs a time step `dt`, a spatial step `dx`, a number of iterations to march forward in time `N`, the coefficient `D`, and an initial function `f` passed as a handle and Dirichlet boundary conditions `alpha` and `beta`. The function should output the solution space matrix `u`. The PDE is given below:

$$v_t = D v_{xx},$$

$$\text{BC}: v(0, t) = \alpha = v(1, t) = \beta = 0;\quad\text{IC}: v(x, 0) = f(x); \quad t > 0, \ x \in [0, 1]$$

<!-- grade_id: cell-2e07bb46c2451706, points: 0, grade: false, solution: true -->
```python
def crankNicolson(dt, dx, N, f, D, alpha, beta):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-347b5f26c968eadc, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
D     = 1
dx    = 0.05
dt    = (dx**2)/(4*D)
alpha = 0
beta  = 0
f     = lambda x: np.exp(-100*(x - 0.5)**2) 
N     = 10
tans  = np.array( [0.          , 0.0000000016, 0.0000001125, 0.0000047851,
                   0.0001234098, 0.0019304541, 0.0183156389, 0.1053992246,
                   0.3678794412, 0.7788007831, 1.          , 0.7788007831,
                   0.3678794412, 0.1053992246, 0.0183156389, 0.0019304541,
                   0.0001234098, 0.0000047851, 0.0000001125, 0.0000000016,
                   0.          ])
ans   = crankNicolson(dt, dx, N, f, D, alpha, beta)
nt.assert_array_almost_equal(tans, ans[0])
print('Test case passed!!!')
```

<!-- grade_id: cell-2f2454e7a0c0acfe, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
D     = 1
dx    = 0.05
dt    = (dx**2)/(4*D)
alpha = 0
beta  = 0
f     = lambda x: np.exp(-100*(x - 0.5)**2) 
N     = 10
tans  = np.array( [0.          , 0.0019064527, 0.0061716818, 0.0167692648,
                   0.0406427189, 0.0876029261, 0.1665959209, 0.277337753 ,
                   0.4013713724, 0.5022251969, 0.5414110658, 0.5022251969,
                   0.4013713724, 0.277337753 , 0.1665959209, 0.0876029261,
                   0.0406427189, 0.0167692648, 0.0061716818, 0.0019064527,
                   0.          ])
ans   = crankNicolson(dt, dx, N, f, D, alpha, beta)
nt.assert_array_almost_equal(tans, ans[-1])
print('Test case passed!!!')
```

<!-- grade_id: cell-7df8b9c44226847f, points: 13, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 2 marks
```