```python
# Initialisation Cell
# You should always put imported modules here
from math import *
import numpy as np
import numpy.testing as nt
import numpy.linalg as LA
from matplotlib import pyplot as plt
import seaborn as sns
#sns.set_style('dark', {'figure.facecolor': '(0,0,0,0)'}) 
sns.set(style="darkgrid")
plt.rcParams['figure.figsize'] = [8, 5]
plt.rcParams['figure.dpi'] = 100
np.set_printoptions(suppress=True, precision=7)
import warnings
warnings.filterwarnings('ignore')
```

# CDES Honours - Lab 4


## Instructions

* Read all the instructions carefully.
* Do not rename the notebook, simply answer the questions and resubmit the file to Moodle.
* **Numpy** has a help file for every function if you get stuck. See: https://docs.scipy.org/doc/numpy-1.15.4/reference/
* See these useful links:
    * https://docs.scipy.org/doc/numpy/user/numpy-for-matlab-users.html
    * https://docs.scipy.org/doc/numpy/user/quickstart.html
* **Numpy** is not always required.
* There are also numerous sources available on the internet, Google is your friend!

## Warmup Exercises

Complete the following problems:

## Question 1

Given a $6 \times 6$ matrix `mat`, for example:
```
1 1 1 0 0 0
0 1 0 0 0 0
1 1 1 0 0 0
0 0 0 0 0 0
0 0 0 0 0 0
0 0 0 0 0 0
```
We can define an hourglass `A` to be a subset of values with indices falling in this pattern in `mat`'s graphical representation, i.e:
```
a b c
  d
e f g
```
Continuing with this example, there are 16 hourglasses of the form above in initial `mat` given. The hourglass sum, is the sum of all of the hourglass's values. Write a function that takes in any matrix `mat` ($\forall\ \mathbb{z}$) and calculates the maximum hourglass sum within `mat`.



<!-- grade_id: cell-70e910fd8da22c87, points: 0, grade: false, solution: true -->
```python
def hourglass(mat):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-9ef45eb345a400c4, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
mat =  [[1, 1, 1, 0, 0, 0],
        [0, 1, 0, 0, 0, 0],
        [1, 1, 1, 0, 0, 0],
        [0, 0, 2, 4, 4, 0],
        [0, 0, 0, 2, 0, 0],
        [0, 0, 1, 2, 4, 0]]

tans = 19
assert(hourglass(mat) == tans)
print('Test case passed!!!')
```

<!-- grade_id: cell-0d867c5cd7f32416, points: 2, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 2 marks
```

<!-- grade_id: cell-f5ba8ee217d424da, points: 2, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 2 marks
```

## Question 2

Write a function that performs `d` left rotations on an array, `a`. For example, if 2 left rotations are applied on the array `[1, 2, 3, 4, 5]`, then the returned array would become `[3, 4, 5, 1, 2]`.

<!-- grade_id: cell-6db4b2c42d0be31b, points: 0, grade: false, solution: true -->
```python
def rotateLeft(a, k):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4809b281086095fa, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
a    = [1, 2, 3, 4, 5]
d    = 4
tans = [5, 1, 2, 3, 4]
assert(rotateLeft(a, d) == tans)
print('Test case passed!!!')
```

<!-- grade_id: cell-eddc930fab9f3b8e, points: 4, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 4 marks
```

## Question 3


Consider a list of integers denoting the positions of people queuing in a line, where the person in the front of the queue is integer 1, the second person 2, all the way to the last person in position $n$. In this queue an individual has the potential to bribe the person directly in front of them and can at most make two brides. Write a function that takes in a array of integers and determines the number of brides required to get the queue into its current state. If the queue is not possible to acheive through the bribing mechanism described above, then the function should return the string `not possible`. 

**Example:**

Given the input `[2, 1, 5, 3, 4]`, the function should return 3. This being the number of bribes to move from `[1, 2, 3, 4, 5]` to the given input.

<!-- grade_id: cell-d10a9f680fac442a, points: 0, grade: false, solution: true -->
```python
def queueBribe(q):
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-a03a7867b4d7926e, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
q = [2, 1, 5, 3, 4]
tans = 3
assert(queueBribe(q) == tans)
print('Test case passed!!!')
```

<!-- grade_id: cell-a2f9d17ca10dff12, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
q = [2, 5, 1, 3, 4]
tans = 'not possible'
assert(queueBribe(q) == tans)
print('Test case passed!!!')
```

<!-- grade_id: cell-8949c5962c51991e, points: 3, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 3 marks
```

## Question 4

A game of tennis consists of two players contests for points within a game. To win a set, a player must be the first to 6 games and win by 2 clear games (if the opposing player is on 5 games, then the winning player will need to win 7 games to win the set). Assuming that the winner of the match, is the player winning the best of 5 sets, write a function that does the following. The function should take in an input string of `A`'s and `B`'s denoting that player `A` or `B` won that particular point. The function should then return the winner of the match, i.e. return `A` or `B`.


**Rules of Tennis:**

* Points are awarded in scores of 15, 30 and 40. 15 represent 1 point, 30 = 2 and 40 = 3. 
* You need 4 points to win a game. If a game lands on 40-40 it’s known as deuce. From deuce a player needs to win 2 consecutive points to win the game. After winning one point from deuce they player is on advantage. If the player wins the next point they win the game, if they lose it goes back to deuce.
* To win the set a player must win 6 games by 2 or more. The opening sets will go to a tie break if its ends up 6-6 where players play first to 7 points. The final set will not have a tie break and requires players to win by two games with no limits.

<!-- grade_id: cell-effbcaf0907eea39, points: 0, grade: false, solution: true -->
```python
def tennisWinner(score):
    # YOUR CODE HERE
    raise NotImplementedError()
    


```

<!-- grade_id: cell-f5e4aa9e23cf20e8, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
score = 'AAAABABBABBBBABAABABBB'
tans = 'B'
assert(tennisWinner(score) == tans)
print('Test case passed!!!')
```

<!-- grade_id: cell-c17a007c936ad082, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
score = 'BBBAAAABABBABBBBAAABABAAAAAABA'
tans = 'A'
assert(tennisWinner(score) == tans)
print('Test case passed!!!')
```

<!-- grade_id: cell-d66b080d5e287463, points: 3, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 3 marks
```

# Main Exercises

## Question 1

Given the equation:
$$
\frac{\partial^{2} T}{\partial r^{2}}+\frac{1}{r} \frac{\partial T}{\partial r}=\frac{1}{4 K} \frac{\partial T}{\partial t}, \quad 0 \leq r \leq 1, \quad t \geq 0,
$$
where:
$$
T^\prime\left(0, t\right)=0, \qquad T(1, t)=100+40 t, \quad T(r, 0)=200r
$$
Write a function that uses central difference approximations for the spatial derivatives and a backward difference approximation for temporal derivatives. The function should take as inputs, the time step `dt`, the spatial step `dr`, a final temperature to approximate `tf`, and the diffusivity coefficient `K`. The function should return the matrix of solutions ended with the approximation at `tf`.

<!-- grade_id: cell-f0a611185c66f60b, points: 0, grade: false, solution: true -->
```python
def heatEqnCylinderNBC(dt, dr, tf, K):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-37001f6c73d5b2a2, points: 2, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 2 marks
# Unit test
dr = 0.01
dt = 0.005
K  = 0.1
tf = 0.5
u = heatEqnCylinderNBC(dt, dr, tf, K)
tans = np.array(  [  0.,   2.,   4.,   6.,   8.,  10.,  12.,  14.,  16.,  18.,  20.,
                    22.,  24.,  26.,  28.,  30.,  32.,  34.,  36.,  38.,  40.,  42.,
                    44.,  46.,  48.,  50.,  52.,  54.,  56.,  58.,  60.,  62.,  64.,
                    66.,  68.,  70.,  72.,  74.,  76.,  78.,  80.,  82.,  84.,  86.,
                    88.,  90.,  92.,  94.,  96.,  98., 100., 102., 104., 106., 108.,
                   110., 112., 114., 116., 118., 120., 122., 124., 126., 128., 130.,
                   132., 134., 136., 138., 140., 142., 144., 146., 148., 150., 152.,
                   154., 156., 158., 160., 162., 164., 166., 168., 170., 172., 174.,
                   176., 178., 180., 182., 184., 186., 188., 190., 192., 194., 196.,
                   198., 100.])
nt.assert_array_almost_equal(u[0, :], tans)
print('Test case passed!!!')
```

<!-- grade_id: cell-379f44bb41763c50, points: 2, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 2 marks
# Unit test
dr = 0.01
dt = 0.005
K  = 0.1
tf = 0.5
u = heatEqnCylinderNBC(dt, dr, tf, K)
tans = np.array([  104.3703931, 104.3730376, 104.3774451, 104.384321 , 104.3937662,
                   104.4058146, 104.4204818, 104.4377766, 104.4577045, 104.4802694,
                   104.5054745, 104.5333223, 104.5638153, 104.5969559, 104.6327464,
                   104.6711895, 104.7122877, 104.7560442, 104.8024621, 104.8515451,
                   104.9032973, 104.9577229, 105.0148271, 105.0746152, 105.1370932,
                   105.2022677, 105.2701459, 105.3407359, 105.4140462, 105.4900863,
                   105.5688665, 105.6503978, 105.7346923, 105.821763 , 105.9116237,
                   106.0042894, 106.0997761, 106.198101 , 106.2992822, 106.4033391,
                   106.5102925, 106.620164 , 106.7329769, 106.8487556, 106.9675257,
                   107.0893144, 107.21415  , 107.3420625, 107.4730831, 107.6072445,
                   107.7445808, 107.8851276, 108.0289219, 108.1760023, 108.3264087,
                   108.4801825, 108.6373668, 108.7980059, 108.9621456, 109.1298333,
                   109.3011176, 109.4760487, 109.6546782, 109.8370589, 110.0232452,
                   110.2132925, 110.4072579, 110.6051992, 110.807176 , 111.0132486,
                   111.2234788, 111.4379292, 111.6566637, 111.8797468, 112.1072444,
                   112.339223 , 112.5757499, 112.8168933, 113.062722 , 113.3133055,
                   113.5687138, 113.8290176, 114.0942876, 114.3645954, 114.6400125,
                   114.9206108, 115.2064623, 115.4976391, 115.7942133, 116.0962567,
                   116.4038414, 116.7170389, 117.0359206, 117.3605573, 117.6910195,
                   118.0273773, 118.3696999, 118.7180559, 119.0725134, 119.4331394,
                   120.       ])
nt.assert_array_almost_equal(u[-1, :], tans)
print('Test case passed!!!')
```

<!-- grade_id: cell-dfd1b5fec14211cf, points: 11, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 11 marks
```

### Question 2

Given the input below. Plot the surface plot of the solution `u` to the problem given in Question 1.

```python
dr = 0.01
dt = 0.005
K  = 0.1
tf = 0.5
u = heatEqnCylinderNBC(dt, dr, tf, K)
```

<!-- grade_id: cell-940ab89c0b356cef, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Question 3

Write a function that solves the 2D Heat equation given below. Specifically, your function should take as inputs, the constant `D`, the stepsizes, `dx`, `dy`, `dt`, as well as the number of time steps to perform `N`. 

$$
v_t = D\left( v_{xx} + v_{yy}\right), 
$$
$$
v(-1, y, t) = v(1, y, t) = v(x, -1, t) = v(x, 1, t) = 0, v(x, y, 0) = e^{-(x^2 + y^2)}  
$$

The function should return the 3D array over all time steps, i.e. $U(x, y, t)$

<!-- grade_id: cell-0b5e0947a573c843, points: 0, grade: false, solution: true -->
```python
def heat2d(dx, dy, dt, D, N):
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-91445c2dc8919927, points: 2, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 2 marks
# Unit test
D  = 1
dx = 0.2
dy = 0.2
dt = dx**2 / 4*D
N  = 10
u = heat2d(dx, dy, dt, D, N)
tans = np.array([[  0.       , 0.       , 0.       , 0.       , 0.       , 0.       ,
                    0.       , 0.       , 0.       , 0.       , 0.       ],
                   [0.       , 0.2780373, 0.3678794, 0.449329 , 0.506617 , 0.5272924,
                    0.506617 , 0.449329 , 0.3678794, 0.2780373, 0.       ],
                   [0.       , 0.3678794, 0.4867523, 0.5945205, 0.67032  , 0.6976763,
                    0.67032  , 0.5945205, 0.4867523, 0.3678794, 0.       ],
                   [0.       , 0.449329 , 0.5945205, 0.726149 , 0.8187308, 0.8521438,
                    0.8187308, 0.726149 , 0.5945205, 0.449329 , 0.       ],
                   [0.       , 0.506617 , 0.67032  , 0.8187308, 0.9231163, 0.9607894,
                    0.9231163, 0.8187308, 0.67032  , 0.506617 , 0.       ],
                   [0.       , 0.5272924, 0.6976763, 0.8521438, 0.9607894, 1.       ,
                    0.9607894, 0.8521438, 0.6976763, 0.5272924, 0.       ],
                   [0.       , 0.506617 , 0.67032  , 0.8187308, 0.9231163, 0.9607894,
                    0.9231163, 0.8187308, 0.67032  , 0.506617 , 0.       ],
                   [0.       , 0.449329 , 0.5945205, 0.726149 , 0.8187308, 0.8521438,
                    0.8187308, 0.726149 , 0.5945205, 0.449329 , 0.       ],
                   [0.       , 0.3678794, 0.4867523, 0.5945205, 0.67032  , 0.6976763,
                    0.67032  , 0.5945205, 0.4867523, 0.3678794, 0.       ],
                   [0.       , 0.2780373, 0.3678794, 0.449329 , 0.506617 , 0.5272924,
                    0.506617 , 0.449329 , 0.3678794, 0.2780373, 0.       ],
                   [0.       , 0.       , 0.       , 0.       , 0.       , 0.       ,
                    0.       , 0.       , 0.       , 0.       , 0.       ]])
nt.assert_array_almost_equal(u[:, :, 0], tans)
print('Test case passed!!!')
```

<!-- grade_id: cell-1f29bdc2bdb458a0, points: 13, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 13 marks
```

### Question 4

Given the input below for Question 3, plot the 2D heat map of the surface over every 20 time steps. Plot each time on its on set of axes.

```python
D  = 1
dx = 0.1
dy = 0.1
dt = dx**2 / 4*D
N  = 100
u  = heat2d(dx, dy, dt, D, N)
```

<!-- grade_id: cell-fe28447e5ea1cd4e, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```