```python
# Initialisation Cell
# You should always put imported modules here
from math import *
import numpy as np
import numpy.testing as nt
import numpy.linalg as LA
import warnings
warnings.filterwarnings("ignore")
from matplotlib import pyplot as plt
np.set_printoptions(suppress=True, precision=7)
```

# APPM4057  - November 2019

# Exam Office Info

| Examiner          | Name              |
|-------------------|-------------------|
| Internal Examiner | Prof B. A. Jacobs |
| External Examiner | Prof N. Hale (Stellenbosch) |

# CDEs Honours - Exam Deferred

## Instructions

* Read all the instructions carefully.
* The exam consists of **80 Marks**, with **three and a half** hours available.
* The written section is to be answered in the book provided.
* You must only access Moodle **TESTS** and NOT Moodle.
* The programming section is to be answered within this Jupyter notebook and resubmitted to Moodle **TESTS**.
* Do not rename the notebook, simply answer the questions and resubmit the file to Moodle.
* The moodle submission link will expire at exactly **12:30** and **NO** late submission will be accepted. Please make sure you submit timeously!
* The **Numpy** and **Matplotlib** documentation has been downloaded and is open on your current machine.
* **NB!!!** Anyone caught using Moodle (and its file storing), flash drives or external notes will be awarded zero and reported to the legal office.

## Written Section

* Answer the following questions in the answer book provided.

### Question 1 - 10 Marks

Discuss each of the below listed terms and their relation to one another (i.e. the theorem which describes this relationship):

* Consistency,
* Stability, 
* Convergence,

Be sure to give potential methods of analysing consistency, stability and convergence of a  linear PDE.

### Question 2 - 10 Marks

Analyse the stability of the given ADI scheme by taking the discrete Fourier Transform of the equation:
$$
u_{j, n}^{m+1 / 2}-u_{j, n}^{m}=\frac{\nu}{\Delta x^{2}} A u_{j, n}^{m+1 / 2}+\frac{\nu}{\Delta y^{2}} B u_{j, n}^{m},
$$
which can be rewritten as:
$$
\left(1-\frac{\beta_{x}}{2} A\right) u^{m+1 / 2}=\left(1+\frac{\beta_{x}}{2} B\right) u^{m+1 / 2},
$$
with $\beta_x = \frac{\Delta t}{\Delta x^2}$ and $\beta_y = \frac{\Delta t}{\Delta y^2}$.

### Question 3 - 10 Marks
Verify that:
$$
\left.\frac{\partial^{3} u}{\partial x^{3}}\right|_{i, j}=\frac{\delta_{+}^{3} u_{i, j}}{(\Delta x)^{3}}+O(\Delta x).
$$
(Hint: Taylor series)

### Question 4 - 10 Marks

Consider the heat equation in cylindrical coordinates:
$$
u_{t}=\frac{\alpha}{r} \frac{\partial}{\partial r}\left(r \frac{\partial u}{\partial r}\right).
$$
(a) [8 Marks] Using a Lax-Wendroff scheme, set up an updating equation.

(b) [2 Marks] Comment on ways in which the scheme could be improved.

## Programming Section 

### Question 1 - 15 Marks

(a) [10 Marks] Consider the wave equation given by:
$$
u_{tt} = c^2u_{xx},
$$
$$
u(x, 0) = x(1 - x), \ \ \ u_t(x, 0) = 0, \ \ \ u(0, t) = u(1, t) = 0.
$$
Write a function that implements a centered space, centered time scheme. Specifically, your function should take as inputs, the stepsizes `dx` and `dt`, the wave speed `c` (not squared valued), the number of iterations to perform `N` and the left and right endpoints `a`, `b`.

Your function should return the solution matrix, containing the wave profile at each time iteration, i.e. `N + 1` rows.

<!-- grade_id: cell-3a4c5c93d86fade5, points: 0, grade: false, solution: true -->
```python
def wave_eq(dx, dt, c, N, ic, a, b):
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-5b37460fed9e3ee6, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
dx = 0.1
c  = 0.2
tf = 8 
N  = 20
dt = (tf - 0)/N
a  = 0 
b  = 1
ic = lambda x: x*(1 - x)
tans = np.array([0.  , 0.09, 0.16, 0.21, 0.24, 0.25, 0.24, 0.21, 0.16, 0.09, 0.  ])
U = wave_eq(dx, dt, c, N, ic, a, b)
nt.assert_array_almost_equal(tans, U[0, :])
print('Test case passed!!!')
```

<!-- grade_id: cell-0f3afb6ecec3131d, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
dx = 0.1
c  = 0.2
tf = 8 
N  = 20
dt = (tf - 0)/N
a  = 0 
b  = 1
ic = lambda x: x*(1 - x)
tans = np.array([0.    , 0.0836, 0.1536, 0.2036, 0.2336, 0.2436, 0.2336, 0.2036,
       0.1536, 0.    , 0.    ])
U = wave_eq(dx, dt, c, N, ic, a, b)
nt.assert_array_almost_equal(tans, U[1, :])
print('Test case passed!!!')
```

<!-- grade_id: cell-b16950b5baa0878c, points: 8, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 8 marks
```

(b) [5 Marks] Continuing with Question 1(a) above, given the input below, produce the 3D surface plot illustrating the evolution of the wave profile over time.

```python
dx    = 0.001
c     = 0.2
tf    = 8 
N     = 2000
dt    = (tf - 0)/N
a     = 0 
b     = 1
ic    = lambda x: x*(1 - x)
xvals = np.linspace(a, b, int((b - a)/dx) + 1)
U     = wave_eq(dx, dt, c, N, ic, a, b)
```

<!-- grade_id: cell-c5c0ad22851f5520, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Question 2 - 15 Marks

Given the equation:
$$
\frac{\partial^{2} T}{\partial r^{2}}+\frac{1}{r} \frac{\partial T}{\partial r}=\frac{1}{4 K} \frac{\partial T}{\partial t}, \quad 0 \leq r \leq 1, \quad t \geq 0,
$$
where:
$$
T^\prime\left(0, t\right)=0, \qquad T(1, t)=100+40 t, \quad T(r, 0)=200\cos(r)
$$


(a) [10 Marks] Write a function which uses central difference approximations for the spatial derivatives and a backward difference approximation for temporal derivatives. The function should take as inputs, the time step `dt`, the spatial step `dr`, a final temperature to approximate `tf`, and the diffusivity coefficient `K`. The function should return the matrix of solutions ended with the approximation at `tf`.

<!-- grade_id: cell-c35fda8283403216, points: 0, grade: false, solution: true -->
```python
def heatEqnCylinderNBC(dt, dr, tf, K):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-636f6cc16e5601dd, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
dr = 0.01
dt = 0.005
K  = 0.1
tf = 0.5
u = heatEqnCylinderNBC(dt, dr, tf, K)
tans = np.array([200.       , 199.9900001, 199.9600013, 199.9100067, 199.8400213,
               199.7500521, 199.640108 , 199.5102001, 199.3603413, 199.1905466,
               199.0008331, 198.7912196, 198.5617272, 198.3123787, 198.0431992,
               197.7542156, 197.4454567, 197.1169534, 196.7687386, 196.400847 ,
               196.0133156, 195.6061829, 195.1794899, 194.733279 , 194.267595 ,
               193.7824843, 193.2779956, 192.7541793, 192.2110877, 191.6487751,
               191.0672978, 190.466714 , 189.8470836, 189.2084687, 188.5509331,
               187.8745426, 187.1793647, 186.4654691, 185.7329271, 184.981812 ,
               184.2121988, 183.4241646, 182.6177881, 181.7931499, 180.9503326,
               180.0894205, 179.2104995, 178.3136576, 177.3989846, 176.4665717,
               175.5165124, 174.5489015, 173.5638359, 172.5614141, 171.5417363,
               170.5049044, 169.4510222, 168.380195 , 167.29253  , 166.1881358,
               165.067123 , 163.9296036, 162.7756913, 161.6055017, 160.4191516,
               159.2167597, 157.9984463, 156.7643332, 155.5145438, 154.249203 ,
               152.9684375, 151.6723752, 150.3611458, 149.0348805, 147.6937117,
               146.3377738, 144.9672021, 143.5821339, 142.1827076, 140.7690631,
               139.3413419, 137.8996866, 136.4442415, 134.975152 , 133.4925652,
               131.9966292, 130.4874936, 128.9653094, 127.4302288, 125.8824053,
               124.3219937, 122.7491499, 121.1640313, 119.5667965, 117.957605 ,
               116.3366179, 114.7039972, 113.0599062, 111.4045094, 109.7379721,
               100.       ])
nt.assert_array_almost_equal(u[0, :], tans)
print('Test case passed!!!')
```

<!-- grade_id: cell-73679076505cf89d, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
dr = 0.01
dt = 0.005
K  = 0.1
tf = 0.5
u = heatEqnCylinderNBC(dt, dr, tf, K)
tans = np.array([199.2194951, 199.1999825, 199.1671443, 199.1157272, 199.0449601,
               198.954583 , 198.8444803, 198.7145953, 198.5649001, 198.3953835,
               198.2060452, 197.996892 , 197.7679365, 197.5191955, 197.2506895,
               196.962442 , 196.6544795, 196.326831 , 195.9795278, 195.6126038,
               195.2260948, 194.8200389, 194.3944763, 193.9494491, 193.4850016,
               193.00118  , 192.4980326, 191.9756095, 191.433963 , 190.873147 ,
               190.2932177, 189.6942329, 189.0762525, 188.4393384, 187.783554 ,
               187.108965 , 186.4156388, 185.7036446, 184.9730535, 184.2239387,
               183.4563747, 182.6704383, 181.8662079, 181.0437636, 180.2031874,
               179.3445631, 178.4679761, 177.5735134, 176.6612638, 175.7313178,
               174.7837671, 173.8187054, 172.8362273, 171.8364292, 170.8194084,
               169.7852636, 168.734094 , 167.6660001, 166.5810825, 165.4794421,
               164.3611798, 163.2263957, 162.0751888, 160.9076564, 159.7238927,
               158.5239885, 157.3080293, 156.0760939, 154.8282524, 153.5645634,
               152.2850713, 150.9898021, 149.6787585, 148.3519142, 147.0092062,
               145.6505255, 144.2757058, 142.8845087, 141.4766063, 140.0515588,
               138.6087868, 137.1475373, 135.6668404, 134.1654569, 132.6418117,
               131.0939118, 129.5192441, 127.914648 , 126.2761573, 124.598803 ,
               122.8763685, 121.1010846, 119.2632497, 117.3507579, 115.3485112,
               113.2376878, 110.994832 , 108.5907209, 105.988954 , 103.1441977,
               100.2      ])
nt.assert_array_almost_equal(u[1, :], tans)
print('Test case passed!!!')
```

<!-- grade_id: cell-105026689c6a6b35, points: 8, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 8 marks
```

(b) [5 Marks] Continuing from Question 2(a) above, plot the surface plot over all time, given the input below:

```python
dr = 0.001
dt = 0.005
K  = 0.1
tf = 0.5
u = heatEqnCylinderNBC(dt, dr, tf, K)
```

<!-- grade_id: cell-790d1930a375743f, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Question 3 - 10 Marks

(a) [7 Marks] Given the heat equation where:
$$
u_t = u_{xx},
$$
$$
0 < x < 1, \ \ \ t \geq 0, \ \ \ u(0, t) = u(1, t) = 0, \ \ \  u(x, 0) = \sin(\pi x),
$$

Implement a Crank-Nicolson scheme on the heat equation given above and find an approximation to $t = 0.5$. The function should take as inputs, the time-step `dt`, the spatial step `dx`, the final time `tf`, and the initial condition as a handle, `f`. The function should return the solution matrix $U$.

Note: the analytical solution is:
$$
u(x, t) = e^{-\pi^2 t}\sin(\pi x)
$$

<!-- grade_id: cell-20c0d37bec86900e, points: 0, grade: false, solution: true -->
```python
def crankNicolson(dt, dx, tf, f):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-032c0e3d4762bc5f, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
dx    = 0.1
dt    = 0.005
tf    = 0.5
f     = lambda x: np.sin(np.pi * x) 
D     = 1
tans  = np.array([0.       , 0.309017 , 0.5877853, 0.809017 , 0.9510565, 1.       ,
                  0.9510565, 0.809017 , 0.5877853, 0.309017 , 0.       ])
ans   = crankNicolson(dt, dx, tf, f)
nt.assert_array_almost_equal(tans, ans[0])
print('Test case passed!!!')
```

<!-- grade_id: cell-cb9befa087fad89c, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
dx    = 0.1
dt    = 0.005
tf    = 0.5
f     = lambda x: np.sin(np.pi * x) 

tans  = np.array([0.       , 0.2942539, 0.5597042, 0.7703667, 0.9056204, 0.9522256,
       0.9056204, 0.7703667, 0.5597042, 0.2942539, 0.       ])
ans   = crankNicolson(dt, dx, tf, f)
nt.assert_array_almost_equal(tans, ans[1])
print('Test case passed!!!')
```

<!-- grade_id: cell-909f04dd96a7776d, points: 5, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 5 marks
```

(b) [3 Marks] Continuing with Question 3(a) above, consider the following input:

```python
dx    = 0.1
dt    = 0.005
tf    = 0.5
f     = lambda x: np.sin(np.pi * x) 
U     = crankNicolson(dt, dx, tf, f)
```

Write a function `errorFunction` which returns the absolute error, relative error and the norm (in that order) at time `tf`, between the true solution given in Question 3(a) and your obtained Crank-Nicolson approximation.

<!-- grade_id: cell-76de265b4f5200db, points: 0, grade: false, solution: true -->
```python
def errorFunction():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1023c1ac7582b912, points: 1, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 1 mark
```

<!-- grade_id: cell-a88a24ba79946105, points: 1, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 1 mark
```

<!-- grade_id: cell-35b4b34fbe37cba6, points: 1, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 1 mark
```